```python
# Initialize Otter
import otter
grader = otter.Notebook("HW1.ipynb")
```

```python
import math
```

**Due Date:** Feb 16, 2026 by 11:59 pm

**Total Points:** 100

**Instructions:**
1. First download this jupyter notebook, (File -> Download) and then upload this jupyter lab (Lab1.ipynb) file on D2L --> HW --> HW1 Dropbox
2.  This lab assignment is due on the date and time specified above. Consult the syllabus for late penalty.
3. This assignment is restricted to individual effort. As per our syllabus, **the use of AI is prohibited**. You may not receive help from any other person except the instructor, the TA, or the AARC (**help from the TA and AARC must be well documented!**).
4. Any resource used (other than the instructor or the course text) must be documented in a new cell of this jupyter notebook detailing the source and describing exactly what was learned and how that information was used in your assignment. Submissions will be severely penalized if copied in part or whole from any source.
5. If you need help, visit your instructor during his posted office hours. If your schedule cannot accommodate any of these times, then email your instructor to schedule a different time..
6. Passing the autograder test does not guarantee that your answer is correct. Be sure to answer questions carefully, as the instructor has additional hidden tests. Full points will only be awarded if your answers pass these tests during grading.
7. If you encounter any difficulties, don't hesitate to contact me for assistance.

# Welcome to HW01.

Learning Objectives:
1. Print statements
2. Errors in Python
3. Python variables and expressions
4. Python functions
5. In-built Data Structures
6. Observation and Causality

# 1. Errors and Comments in Python
In this section, your task is to locate the errors in the Python code and correct them through debugging. A helpful debugging strategy is to comment out all the lines of code first, then uncomment and execute the program one line at a time. If an error appears, fix that line before proceeding to the next one.

<!-- BEGIN QUESTION -->

### Question 1.1 [5 points]

Identify and correct the error in the code below.

The program should display the following message verbatim:<br>
`Welcome to CSCI1462`

Add your corrected code in a new cell.

Note: There shouldn't be any extra space in the printed message. 

```python
msg1 = "Welcome to "
msg2 = CSCI
msg3 = 1462

print(f"{msg1" {msg2}       {msg3}")
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1.2 [5 points]

Identify and correct the error in the code below.

The code is trying to compute the mathematical expression:

$$
10 \cdot \log_2 \Biggl( \left( 1 + \frac{5}{6 \cdot 2} \right)^{3 \cdot 0.5} \Biggr)
$$

Add your corrected code in a new cell.

Note: Make sure to use the PEMDAS rule correctly. 

See [math.log](https://www.w3schools.com/python/ref_math_log.asp) function to calculate the log

```python
expr = 10 * log(1 + 5 / 6 * 2, 2) ^ 3 * 0.5
```

```python
expr
```

```python
grader.check("q1.2")
```

## 2. Variables

<!-- BEGIN QUESTION -->

### Question 2.1 [5 points]

For the sequence of code provided below, indicate the values stored in the variables after each statement is executed (i.e., although this is one program, we are executing each statement in separate cells to observe the current values of the variables). If a variable does not have a value at that point, write "N/A"

The first example has been completed for you as a reference: 

| Line number | Code | Value of x | Value of y | Value of z |
| ----- | ----- | :----: | :-----: | :----: | 
| (1) |  x = 1 | 1 | N/A | N/A |
| (2) | y = 1 | | | |
| (3) | z = 3 | | | |
| (4) | x = z | | | |
| (5) | z = 5 | | | |
| (6) | x = y + 2 | | | |
| (7) | y = y - 3 | | | |
| (8) | z = x + y | | | |

Note: To verify your solution, execute each line of code in separate cells and use `print(f"x: {x}, y: {y}, z: {z}")` to observe the values of the variables after executing each line.

Format your answer as follows: `(1) x: 1, y: N/A, z: N/A`

Here, `1)` indicates the line number, `x: 1` shows the value of x after line 1, `y: N/A` indicates that y has no value yet after line 1, and `z: N/A` shows that z also has no value after line 1.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.2 [5 points]

Evaluate the following numerical expressions in your head, and write your answer in the markdown cell below. 

(1) 5 ** 2 

(2) 9 * 5

(3) 15 / 12

(4) 12 / 15

(5) 15 // 12

(6) 12 // 15

(7) 5 % 2

(8) 9 % 5

(9) 15 % 12

(10) 12 % 15

(11) 6 % 6

(12) 0 % 7

Note: You may create a code cell to check your results.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2.3 [5 points]

Newton's law of universal gravitation states that every particle attracts every other particle in the universe with a force (in Newton) that is given as follows:
$$F = \frac{G * m_1 * m_2}{d^2}$$
Where G is the gravitational constant, $m_1$ and $m_2$ are the mass of the two objects, and d is the distance between them. 

Complete the following code that computes the gravitational force (in Newton) exerted by the moon on my friend, *John* who is 160 pounds. Store your answer in the variable `F`, rounded upto two decimal points.

Hint: 1 pound is approximately equal to 0.453 kilograms.

Given $G = 6.674 * 10^{-11} Nm^2{kg}^2$, <br>
mass of moon $ = 6 * 10^{24} kg$, <br>
distance between moon and earth's surface is $3.84 * 10^8 m$

```python
G = ...
moon_mass = ...
john_mass = ...
d = ...

F = ...
F
```

```python
grader.check("q2.3")
```

### Question 2.4 [5 points]

Many people keep time using a 24 hour clock (11 is 11am and 23 is 11pm, 0 is midnight). If it is currently 13 and you set your alarm to go off in 50 hours, it will be 15 (3pm). Complete the Python program to solve the general version of the above problem. The variable `curr_time` stores the time now (in hours), and the variable `num_hrs` stores the number of hours to wait for the alarm. Your program should calculate what time will be on the clock when the alarm goes off. Store your answer in the variable `set_time`.

Note: Try more test cases by changing the values of `curr_time` and `num_hrs` to validate if your solution is correct. 

```python
curr_time = 13
num_hrs = 50

set_time = ...
set_time
```

```python
grader.check("q2.4")
```

# 3. Functions

<!-- BEGIN QUESTION -->

### Question 3.1 [5 points]
Strings and numbers are distinct data types. For instance, 5 is a number (integer type), while "5" is a string (string type).

Attempting to add a number (int) and a string (str), such as 5 + "5", results in an error because the operation is not valid.

Fortunately, Python provides [built-in functions](https://docs.python.org/3/library/functions.html) to convert between these data types: you can convert a number to a string and vice versa.

|Function name | Description | Example |
| --- | --- | --- |
| int | convert a string of digits or a float (real number) to an integer (int) value | int("123") would result in 123 of int type |
| float | convert a string of digits (perhaps with decimal points) or an int to a float (real number) value | float("123") would result in 123.0 of float type |
| str | convert any number (integer or real number) to a string | str(123) would result in `123` of string type |

What is the value of `example` after executing the following code?

```
expr = 2 + int("10") + float("8") # When a float and int is added, the result is float.
example = "Salary" + " " + str(expr)
```

Please provide the answer directly in the markdown cell below, without running the code.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3.2 [20 points (2 points each)]

Consider the following function named `greet_user` that takes the following parameters:

- name (string) – the name of the user to greet.
- greeting (string, default value: "Hello") – the greeting message.
- punctuation (string, default value: "!") – the punctuation mark to end the greeting.
    
The function return a greeting string that combines these input parameters in the following format:

```greeting, name punctuation```

```python
def greet_user(name, greeting="Hello", punctuation="!"):
    return f"{greeting}, {name} {punctuation}"
```

Determine the output of each of the following function calls. Write your answer directly without executing the code. If a call results in an error, write `Error`. Also, give an explanation for why it is throwing an error.

1. greet_user("John", "Hi", ":)")
2. greet_user(John, 'Hi', ":)")
3. greet_user("John")
4. greet_user("John", punctuation="!!!")
5. greet_user("John", "---")
6. greet_user(greeting="Hey", punctuation="!")
7. greet_user(greeting="Hey", punctuation="!", "John")
8. greet_user()
9. greet_user("John", greeting="Hola", "!!!")
10. greet_user("John", greeting="Hola", punctuation="!!!")

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 4. In-built Data Structures

### Question 4.1 [5 points]

Utilize the [replace](https://www.w3schools.com/python/ref_string_replace.asp) string method to convert the text `HitchHiker` into `MatchMaker`. Note that you should replace only specific substrings within the original string, rather than replacing the entire string.

Store your result in the variable, `your_ans`. 

Do NOT replace the entire string with another string. In other words, do not replace entire `HitchHiker` string with `MatchMaker` string. Simply replace substring within `HitchHiker` to make it `MatchMaker`.  You only need to replace the substring `Hi`.

```python
my_str = "HitchHiker"
your_ans = ...
your_ans
```

```python
grader.check("q4.1")
```

### Question 4.2 [25 points (5 points each)]

Consider the following tweet text

`"Amazing!!! Visit https://example.com NOW #AI #MachineLearning 🤖🔥"`

Complete the python program that cleans the text using only Python [string methods](https://www.w3schools.com/python/python_ref_string.asp). 

(a) Convert the tweet to lowercase

(b) Remove the URL

(c) Remove the # symbol but keep the hashtag words

(d) Remove emojis from the end of the text

(e) Remove extra spaces

Your output should produce the text `amazing!!! visit now ai machinelearning`

```python
tweet = "Amazing!!! Visit https://example.com NOW #AI #MachineLearning 🤖🔥"
```

### Question 4.2(a)

Convert the tweet to lowercase. 

```python
tweet = ...
tweet
```

```python
grader.check("q4.2(a)")
```

### Question 4.2(b)

Remove the URL from the `tweet` variable after you solved Q4.2(a). 

Hint: Use `replace` method to replace the URL with empty string `""`

```python
tweet = ...
tweet
```

```python
grader.check("q4.2(b)")
```

### Question 4.2(c)

Remove the `#` symbol from the `tweet` variable, but keep the hashtag words, after you have solved Q4.2(b).

```python
tweet = ...
tweet
```

```python
grader.check("q4.2(c)")
```

### Question 4.2(d)

Remove emojis from the `tweet` variable after you have solved Q4.2(c)

Hint: Use [str.rstrip()](https://www.w3schools.com/python/ref_string_rstrip.asp) method

```python
tweet = ...
tweet
```

```python
grader.check("q4.2(d)")
```

### Question 4.2(e)

Finally, remove extra spaces from the `tweet` variable.

Hint: 
1. Use `str.split()` method to first split each word of strings into a list
2. Use `str.join()` method to join each of those words with a space in between

```python
tweet = ...
tweet
```

```python
grader.check("q4.2(e)")
```

### Question 4.3 [5 points]

Write a Python code to extract the value 600 in the list provided below.

Store the output in the variable `your_ans`.
```
[1, 2, [30, 40, [500, 600, 700], 40, 30], 2, 1]
```

```python
my_list = [1, 2, [30, 40, [500, 600, 700], 40, 30], 2, 1]
your_ans = ...
your_ans
```

```python
grader.check("q4.3")
```

# 5. Observation vs Causality

<!-- BEGIN QUESTION -->

### Question 5.1 [5 points]

The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong. (The history of the event is detailed [here](https://pmc.ncbi.nlm.nih.gov/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? Select the correct choice from the following. Give reasoning behind why you picked one option over the others. 

1. If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2. Survivors could provide additional information about what else could have **caused** the cholera, potentially unearthing another **cause**.

3. Through considering the survivors, Whitehead could have identified a cure for cholera.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 5.2 [5 points]

To establish the causality between literacy and health relationship, which of the following would be the best way to test? Select only one answer.

1. Run a randomized controlled experiment

2. Brainstorm some confounding factors, and check if any of them have association with literacy and health. 

3. Ask people in the poverty striken neighborhood to find how many of them are illiterate and unhealthy, and if the number is greater, then we can conclude the causality. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q1.2 (points: None)
```
>>> math.isclose(round(expr, 2), 7.54) == True
True
```


Tests for q2.3 (points: None)
```
>>> math.isclose(F, 0.2) == True
True
```


Tests for q2.4 (points: None)
```
>>> set_time == 15
True
```


Tests for q4.1 (points: None)
```
>>> your_ans == 'MatchMaker'
True
```


Tests for q4.2(a) (points: None)
```
>>> len(tweet) == 74
False
```


Tests for q4.2(b) (points: None)
```
>>> ('https' in tweet) == False
True
```


Tests for q4.2(c) (points: None)
```
>>> ('#' in tweet) == False
True
```


Tests for q4.2(d) (points: None)
```
>>> ('🔥' in tweet) == False
True
```
```
>>> ('🤖' in tweet) == False
True
```


Tests for q4.2(e) (points: None)
```
>>> len(tweet) == 39
True
```


Tests for q4.3 (points: None)
```
>>> your_ans == 600
True
```
