```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab1.ipynb")
```

```python
import numpy as np
import math
```

**Due Date:** Jan 29, 2026 by 11:59 pm

**Total Points:** 100

**Instructions:**
1. First download this jupyter notebook, (File -> Download) and then upload this jupyter lab (Lab1.ipynb) file on D2L --> Lab --> Lab1_Dropbox
2.  This lab assignment is due on the date and time specified above. Consult the syllabus for late penalty.
3. This assignment is restricted to individual effort. As per our syllabus, **the use of AI is prohibited**. You may not receive help from any other person except the instructor, the TA, or the AARC (**help from the TA and AARC must be well documented!**).
4. Any resource used (other than the instructor or the course text) must be documented in a new cell of this jupyter notebook detailing the source and describing exactly what was learned and how that information was used in your assignment. Submissions will be severely penalized if copied in part or whole from any source.
5. If you need help, visit your instructor during his posted office hours. If your schedule cannot accommodate any of these times, then email your instructor to schedule a different time..
6. Passing the autograder test does not guarantee that your answer is correct. Be sure to answer questions carefully, as the instructor has additional hidden tests. Full points will only be awarded if your answers pass these tests during grading.
7. If you encounter any difficulties, don't hesitate to contact me for assistance.

# Welcome to Lab01.

Learning Objectives:
1. Print statements
2. Errors in Python
3. Python variables and expressions
4. Python functions

## 1. Errors and Comments in Python

In this section, your task is to locate the errors in the Python code and correct them through debugging. A helpful debugging strategy is to comment out all the lines of code first, then uncomment and execute the program one line at a time. If an error appears, fix that line before proceeding to the next one.

### Question 1.1 [5 points]:

Identify and correct the error in the code below.

The program is intended to print the following message:
`I love CSCI 1462: Data Analytics. But the final exam is at 7:30 am in the morning, so I need to set an alarm for it so that I don't miss it.`

Add your corrected code in a new cell.

```python
msg1 = "I love CSCI 1462: Data Analytics.

msg2 = But the final exam is at 7:30 am in the morning, so I need to set an alarm for it so that I don't miss it.

print(f"{msg1} {msg2}")
```

Output:
```
Cell In[124], line 1
    msg1 = "I love CSCI 1462: Data Analytics.
           ^
SyntaxError: unterminated string literal (detected at line 1)
```

### Question 1.2 [5 points]:

Identify and correct the error in the code below.

The program should display the following message:
`800 seconds is equivalent to 13 minutes and 20 seconds.`

Add your corrected code in a new cell.

```python
seconds = "800"
mins = seconds // 60
remaining_secs = secs % 60

print(f"{secs} seconds is equivalent to {mins} minutes and {remaining_secs} seconds.")
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[125], line 2
      1 seconds = "800"
----> 2 mins = seconds // 60
      3 remaining_secs = secs % 60
      5 print(f"{secs} seconds is equivalent to {mins} minutes and {remaining_secs} seconds.")

TypeError: unsupported operand type(s) for //: 'str' and 'int'
```

### Question 1.3 [5 points]:

Locate and correct the error in the code below.

The program should print the following message:
`The total income is 82000 USD`

Add your corrected code in a new cell.

```python
bonus_income = "2500"
salary = 80,000 

total_income = salary + bonus_income
print(f"The total income is {total_income} USD")
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[126], line 4
      1 bonus_income = "2500"
      2 salary = 80,000 
----> 4 total_income = salary + bonus_income
      5 print(f"The total income is {total_income} USD")

TypeError: can only concatenate tuple (not "str") to tuple
```

## 2. Variables and Expressions

### Question 2.1 [10 points]:

You can assign numbers 0 through 6 to the days of the week, where 0 represents Sunday and 6 represents Saturday. For example, if you leave for a holiday on day 3 (Wednesday) and stay for 10 nights, you would return on day 6 (Saturday).

Complete the following general version of the program which accepts a value (between 0 to 6) for the starting day number, and the length of your stay (any non-negative number), and it will tell you the number of day of the week you will return on.

Hint: 
- Use `modulo (%)` operator

```python
starting_day_number = 3 # Try to change this number and check if the output is correct with this new change
length_of_stay = 10 # Try to change this number and check if the output is correct with this new change

day_of_return = (starting_day_number + length_of_stay) % 7 # SOLUTION

print(f"You will return on the day number {day_of_return}")
```

Output:
```
You will return on the day number 6
```

```python
grader.check("q2.1")
```

### Question 2.2 [5 points]:

The formula for computing the final amount if one is earning compound interest is given on Wikipedia as

$$
A = P (1 + \frac{r}{n})^{nt}
$$

where, 
- P = principal amount (initial investment)
- r = annual nominal interest rate (as a decimal)
- n = number of times the interest is compounded per year
- t = number of years


Write a Python program that assigns the principal amount of 10000 to variable **P**, assign to **n** the value 12, assign to **t** the value of 4, and assign to **r** the interest rate of 8% (0.08). Calculate and print the final amount after **t** years (_rounded upto two decimal points_). Your program should print out the message in the format, `The final amount after 4 years is 10269.35`

```python
P = 10000 # SOLUTION
n = 12 # SOLUTION
r = 0.08 # SOLUTION
t = 4 # SOLUTION

A = round(P * (1 + r/n)**t, 2) # SOLUTION


"""# BEGIN PROMPT
print(f"The final amount after {...} years is {...}")
"""# END PROMPT
print(f"The final amount after {t} years is {A}") # SOLUTION NO PROMPT
```

Output:
```
The final amount after 4 years is 10269.35
```

```python
grader.check("q2.2")
```

### Question 2.3 [10 points]:

Write a Python expression in the next cell that evaluates the following expression and save the output in a variable named _result_:

$3∗(4\frac{3}{5})+e^{0.5∗\pi}−\frac{11}{36}+ \sqrt{2\pi}$

Please note that by $a\frac{b}{c}$, I mean $a * \frac{b}{c}$ 

Make sure to use the predefined `pi` and `e` constants from Python’s math library. 

```python
result = 3 * (4 * 3/5) + math.e **(0.5 * math.pi) - 11/36 + math.sqrt(2 * math.pi) # SOLUTION
result
```

Output:
```
14.211550100040796
```

```python
grader.check("q2.3")
```

### Question 2.4 [5 points]:

Take the sentence: _`All work and no play makes Jack a dull boy`_. Store each word in a separate variable, then print out the sentence on one line using print. You may want to read the Q1.1 of the Lab01_Practice Worksheet to know how to print statements.

```python
var1 = "All" # SOLUTION
var2 = "work" # SOLUTION
var3 = "and" # SOLUTION
var4 = "no" # SOLUTION
var5 = "play" # SOLUTION
var6 = "makes" # SOLUTION
var7 = "Jack" # SOLUTION
var8 = "a" # SOLUTION
var9 = "dull" # SOLUTION
var10 = "boy" # SOLUTION

"""# BEGIN PROMPT
print(f"...")
"""# END PROMPT

print(f"{var1} {var2} {var3} {var4} {var5} {var6} {var7} {var8} {var9} {var10}") # SOLUTION NO PROMPT
```

Output:
```
All work and no play makes Jack a dull boy
```

## 3. Calling a Function

### Question 3.1 [5 points]:

Take a look at the following function `repeatMsg()`, which accepts two parameters: a `str` and an `int`. The function repeats the string the specified number of times.

For example, `repeatMsg("hey", 2)` would return `heyhey`, showing that the message `hey` is repeated twice.

```python
# You may ignore this cell, since we do not know the body of the function to answer the questions. 
def repeatMsg(msg, ntimes):
    return msg * ntimes
```

```python
repeatMsg("hey", 2)
```

Output:
```
'heyhey'
```

#### Task: 
Call the `repeatMsg` function correctly such that it returns the message, `hellohellohello`. Store this result in the variable, `your_ans`.

```python
your_ans = repeatMsg("hello", 3) # SOLUTION
print(your_ans)
```

Output:
```
hellohellohello
```

```python
grader.check("q3.1")
```

<!-- BEGIN QUESTION -->

### Question 3.2 [5 points]:

What happens if I call the function like the following: `repeatMsg(3, "hello")`? Will the function return an error? If so, why?

_Type your answer here, replacing this text._

The function won't return an error. This is because 3 is passed as the first argument to the parameter msg and "hello" is passed as the second argument to the parameter ntimes, as a result the function body would do: 3 * "hello", which will print the message hello three times.

<!-- END QUESTION -->

### Question 3.3 [5 points]:

Let's call the function using keyword arguments such that it outputs, `hellohellohello`

```python
"""# BEGIN PROMPT
your_ans = repeatMsg( ... = ..., ... = ...)
"""# END PROMPT
your_ans = repeatMsg(msg="hello", ntimes=3) # SOLUTION NO PROMPT
your_ans
```

Output:
```
'hellohellohello'
```

```python
grader.check("q3.3")
```

<!-- BEGIN QUESTION -->

### Question 3.4 [40 points (5 points for each part)]:

Now suppose we modify the function, `repeatMsg()` so that it has default values: `"hello"` for the first parameter (`msg`) and `2` for the second parameter (`ntimes`), as shown below:

```python
def repeatMsg(msg="hello", ntimes=2):
    return msg * ntimes
```

Determine the output of each of the following function calls. Write your answer directly without executing the code. If a call results in an error, write `Error`

1. repeatMsg("Hola", 3)
2. repeatMsg(4)
3. repeatMsg("Hola")
4. repeatMsg()
5. repeatMsg(msg="Hi", 3)
6. repeatMsg(msg="Hi", ntimes=3)
7. repeatMsg(3, 5)
8. repeatMsg("3", 3)

_Type your answer here, replacing this text._

1. `HolaHolaHola`
2. `8`
3. `HolaHola`
4. `hellohello`
5. `Error`
6. `HiHiHi`
7. `15`
8. `333`

<!-- END QUESTION -->



Tests for q2.1 (points: None)
```
>>> 0 <= day_of_return <= 6
True
```


Tests for q2.2 (points: None)
```
>>> math.isclose(A, 10269.35) == True
True
```


Tests for q2.3 (points: None)
```
>>> math.isclose(result, 14.21, rel_tol=0.01)
True
```


Tests for q3.1 (points: None)
```
>>> your_ans == 'hellohellohello'
True
```


Tests for q3.3 (points: None)
```
>>> your_ans == 'hellohellohello'
True
```
